[![Open In Colab](https://colab.research.google.com/assets/colab-badge.svg)](https://colab.research.google.com/github/upm-classes/image-understanding-2022-2023/blob/main/practice1/practice_1.ipynb)

# Practice 1

This is the first practice of image understanding in which we are going to code some of the things we have already seen during classes. 

Fill with your data:

- Full name 1: 
- Full name 2: 

**DO NOT MODIFY (ADD, DELETE,...) ANY CELL. YOU HAVE TO WRITE YOUR CODE ONLY IN THE SPECIFIED CELLS.**

```python
# These are the libraries that should be used. 
# NO OTHER LIBRARIES ARE ALLOWED
from typing import Tuple
from skimage.io import imread
from skimage.transform import resize
import numpy as np
from scipy.signal import convolve2d
import matplotlib.pyplot as plt # Only to plot the images
```

```python
# Reading image
image = imread('kittycat.jpg')

# Plotting image
plt.figure(figsize=(2,2))
plt.imshow(image)
plt.show()
```

## Task 1: Basic image statistics (2 points)

<!-- grade_id: image_statistics, points: 0, grade: false, solution: true -->
```python
def image_statistics(image: np.array) -> Tuple[float, float, float]:
    """
    Given a grayscale image of 8-bits, it should return a tuple with
    the maximum, minimum and average values of the
    pixels.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: image_statistics_test, points: 1, grade: true, solution: false -->
```python
# Tests that your function must pass

assert image_statistics(image[:,:,0])[0] == 255
assert image_statistics(image[:,:,0])[1] == 0
assert np.isclose(image_statistics(image[:,:,0])[2], 131.103)

```

<!-- grade_id: histogram, points: 0, grade: false, solution: true -->
```python
def calculate_histogram(image: np.array, bins: int = 255) -> Tuple[np.array, np.array]:
    """
    Given a grayscale image of 8-bits, it should return two numpy 
    arrays: the first one a numpy array with the histogram
    and the second one with the bin edges according to the 
    number of bins selected. You should use a numpy 
    function to calculate the histogram.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

To check if your answer is correct you should obtain 
the same graphic as this one in the next cell.

![image](hist.png)


```python
hist, edges = calculate_histogram(image[:,:,0])

plt.figure(figsize=(4,3))
plt.bar(edges[1:], hist)
plt.xlabel('Pixel values')
plt.ylabel('Frequency')
plt.show()
```

<!-- grade_id: histogram_test, points: 1, grade: true, solution: false -->
```python
try:
    assert (calculate_histogram(image[:,:,0], 5)[0] == np.array([8331,54144,55237,33667,25021])).all()
except:
    assert (calculate_histogram(image[:,:,0], 5)[0] == np.array([8340, 54234, 55086, 33685, 25055])).all()
```

### Task 2: Image transformations (4 points)

<!-- grade_id: invert_image, points: 0, grade: false, solution: true -->
```python
def invert_image(image: np.array) -> np.array:
    """
    Function that inverts a grayscale image of 8-bits. Dark areas should be brighter 
    and bright areas should be darker.It is equal to the operation (255 - image).
    The result should be a uint8 numpy array.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
fig, axes = plt.subplots(nrows=1, ncols=2, sharex=True, sharey=True, figsize=(5, 2))
im_0 = axes[0].imshow(image[:,:,0], cmap='gray')
fig.colorbar(im_0, ax=axes[0])
im_1 = axes[1].imshow(invert_image(image[:,:,0]), cmap='gray')
fig.colorbar(im_1, ax=axes[1])
plt.show()
```

<!-- grade_id: invert_image_test, points: 1, grade: true, solution: false -->
```python
assert invert_image(image[:,:,0]).dtype == 'uint8'
assert (invert_image(np.array([233, 255, 0], dtype='uint8')) == np.array([22,0,255], dtype='uint8')).all()
```

<!-- grade_id: rotate_flip_image, points: 0, grade: false, solution: true -->
```python
def rotate_flip_image(image: np.array, rotation: int = 0, flip: bool = False) -> np.array:
    """
    This function should rotate an rgb image 0, 90, 180 or 270 degrees. 
    Other angle should return the following message: 'Error: the angle specified is not valid.'.  
    When flip is True, the image is mirrored over the x axis. 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Visualization of the results

fig, axes = plt.subplots(nrows=2, ncols=2, figsize=(4, 4))
im_0 = axes[0,0].imshow(image)
im_1 = axes[0,1].imshow(rotate_flip_image(image, rotation=90))
im_2 = axes[1,0].imshow(rotate_flip_image(image, rotation=90, flip=True))
im_3 = axes[1,1].imshow(rotate_flip_image(image, rotation=270, flip=True))
plt.setp(axes, xticks=[], yticks=[])
plt.show()
```

<!-- grade_id: rotate_flip_image_test, points: 1, grade: true, solution: false -->
```python
assert rotate_flip_image(image, rotation=90, flip=True).dtype == 'uint8'
assert rotate_flip_image(image, rotation=1) == 'Error: the angle specified is not valid.'
assert (rotate_flip_image(np.array([[1,2],[3,4]],dtype='uint8'), rotation=90, flip=True) == np.array([[4,2],[3, 1]], dtype='uint8')).all()
assert (rotate_flip_image(np.array([[1,2],[3,4]],dtype='uint8'), rotation=180) == np.array([[4,3],[2, 1]], dtype='uint8')).all()
assert (rotate_flip_image(np.array([[1,2],[3,4]],dtype='uint8'), rotation=270, flip=True) == np.array([[1,3],[2, 4]], dtype='uint8')).all()
```

<!-- grade_id: color_depth, points: 0, grade: false, solution: true -->
```python
def color_depth(image: np.array, number_bits: int) -> np.array:
    """
    The input is a grayscale image of 8-bits. The output image should be rescaled to 
    the number of bits indicated. The data type of the output should be uint16.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Visualization of the results

fig, axes = plt.subplots(nrows=2, ncols=2, sharex=True, sharey=True, figsize=(6, 4))
number_bits = 8
im_0 = axes[0,0].imshow(color_depth(image[:,:,0], number_bits), vmax=2**number_bits, vmin=0, cmap='gray')
fig.colorbar(im_0, ax=axes[0,0])
number_bits = 4
im_1 = axes[0,1].imshow(color_depth(image[:,:,0], number_bits), vmax=2**number_bits, vmin=0, cmap='gray')
fig.colorbar(im_1, ax=axes[0,1])
number_bits = 2
im_2 = axes[1,0].imshow(color_depth(image[:,:,0], number_bits), vmax=2**number_bits, vmin=0, cmap='gray')
fig.colorbar(im_2, ax=axes[1,0])
number_bits = 1
im_3 = axes[1,1].imshow(color_depth(image[:,:,0], number_bits), vmax=2**number_bits, vmin=0, cmap='gray')
fig.colorbar(im_3, ax=axes[1,1])
plt.setp(axes, xticks=[], yticks=[])
plt.show()
```

<!-- grade_id: color_depth_test, points: 1, grade: true, solution: false -->
```python
assert color_depth(np.array([0, 125, 255], dtype='uint8'), 8).dtype == 'uint16'
assert (color_depth(np.array([0, 125, 255], dtype='uint8'), 8) == np.array([0, 125, 255], dtype='uint16')).all()
assert (color_depth(np.array([0, 125, 255], dtype='uint8'), 4) == np.array([0, 7, 15], dtype='uint16')).all()
assert (color_depth(np.array([0, 125, 255], dtype='uint8'), 2) == np.array([0, 1, 3], dtype='uint16')).all()
```

<!-- grade_id: halve_image, points: 0, grade: false, solution: true -->
```python
def halve_image(image: np.array, times: int) -> np.array:
    """
    The input image is an 8-bit rgb image. The output image must be reduced by 
    half the number of times indicated. If the new size is not an integer you  
    have to truncate the decimals. The output image should be an 8-bit image.
    You can use the resize function of scikit-image with preserve_range=True,
    No open-cv function is allowed.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Visualization of the results

fig, axes = plt.subplots(nrows=2, ncols=2, figsize=(6, 4))
im_0 = axes[0,0].imshow(halve_image(image, times=1))
im_1 = axes[0,1].imshow(halve_image(image, times=2))
im_2 = axes[1,0].imshow(halve_image(image, times=3))
im_3 = axes[1,1].imshow(halve_image(image, times=4))
plt.show()
```

<!-- grade_id: halve_image_test, points: 1, grade: true, solution: false -->
```python
assert halve_image(image, times=2).dtype == 'uint8'
assert (halve_image(image, 0).shape == np.array([420, 420, 3])).all()
assert (halve_image(image, 3).shape == np.array([52, 52, 3])).all()
assert (halve_image(image, 5).shape == np.array([13, 13, 3])).all()
assert (halve_image(image, 8) == np.array([[[131, 149,  99]]], dtype='uint8')).all()

```

### Spatial filters (2 points)

```python
def convolve(image: np.array, kernel: np.array) -> np.array:
    """
    Funtion to convolve a grayscale image given a kernel.
    """
    im = image.astype(np.float32)
    result = convolve2d(im, kernel, boundary='symm', mode='same')
    return result
```

```python
# Example of the effect of a kernel for smoothing an image
kernel = np.array([[1,1,1],[1,9,1], [1,1,1]])

image_channel = image[:,:,0]
convolved = convolve(image_channel, kernel)

fig, axes = plt.subplots(nrows=3, ncols=2, figsize=(6, 8))
im_0 = axes[0,0].imshow(image_channel, cmap='gray')
fig.colorbar(im_0, ax=axes[0,0])
im_1 = axes[0,1].imshow(convolved, cmap='gray')
fig.colorbar(im_1, ax=axes[0,1])
im_2 = axes[1,0].imshow(image_channel[50:100, 50:100], cmap='gray')
fig.colorbar(im_2, ax=axes[1,0])
im_3 = axes[1,1].imshow(convolved[50:100, 50:100], cmap='gray')
fig.colorbar(im_3, ax=axes[1,1])
im_4 = axes[2,0].imshow(image_channel[100:150, 100:150], cmap='gray')
fig.colorbar(im_4, ax=axes[2,0])
im_5 = axes[2,1].imshow(convolved[100:150, 100:150], cmap='gray')
fig.colorbar(im_5, ax=axes[2,1])
plt.setp(axes, xticks=[], yticks=[])
plt.show()
```

<!-- grade_id: sobel, points: 0, grade: false, solution: true -->
```python
def sobel_filtering(image: np.array) -> np.array:
    """
    Define the function using the filters sobel_x and sobel_y using function "convolve" 
    defined before. The input is an grayscale image of 8-bits. The output is a gradient
    image calcluated using sobel operators both x and y.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Visualization of the results

filtered_image = sobel_filtering(image[:,:,0])

fig, axes = plt.subplots(nrows=1, ncols=2, figsize=(6, 8))
im_0 = axes[0].imshow(image_channel, cmap='gray')
im_1 = axes[1].imshow(filtered_image, cmap='gray')
plt.setp(axes, xticks=[], yticks=[])
plt.show()
```

<!-- grade_id: sobel_test, points: 2, grade: true, solution: false -->
```python
sample = np.array([[149.271,117.889,51.478],[53.907,90.785,62.6258],
                   [63.529,43.863,93.648]])
assert np.isclose(sobel_filtering(image[:,:,0])[100:103, 99:102], sample).all()
```